Inainte sa trimiti notebook-ul:

* asigura-te ca totul se executa cum trebuie. Prima oara, **restarteaza kernel-ul** (in meniu, selecteaza Kernel$\rightarrow$Restart & Run all);

* nu lasa niciun mesaj la output (`disp`, `printf`, linii care nu sunt incheiate cu `;`) in functiile pe care le scrii decat daca ti se specifica explicit sa faci acest lucru! Orice abatere va fi **depunctata automat** de grader;

* citeste instructiunile de [aici](https://github.com/radusqrt/numerical-methods/blob/master/lab_instructions.md) cu atentie;

* asigura-te ca rezolvi toate celulele in care scrie `YOUR CODE HERE` sau `YOUR ANSWER HERE` si ca iti treci numele mai jos:

```octave
STUDENT_NAME = "PETRUC RARES";
```

---

# Laborator 7-8: Aflarea valorilor si vectorilor proprii
### Autor: Radu Stochitoiu

## Problema 1

Implementati funcția ```power_method(A, x0, tol, max_steps)``` ce primește o matrice pătratică $A$, o aproximare ințială a vectorului propriu domninant $x_0$, o toleranță sub care se poate opri algoritmul $tol$ și un număr maxim de pași $max\_steps$ și returnează o valoarea proprie dominantă a matricei $A$, $lambda$, un vector propriu normat asociat valorii proprii $lambda$, $x$, eroarea de la ultimul pas $err$ și numărul de pași efectuați $steps$.

Eroarea de la fiecare pas se calculează în funcție de vectorul propriu aproximat, nu de valoarea proprie aproximată.

<!-- grade_id: cell-82eacb35e4df3e68, points: 0, grade: false, solution: true -->
```octave
function [lambda, x, err, steps] = power_method(A, x0, tol, max_steps)

    steps = 0;
    prev_x = x0;
    while 1
    
        x = A*prev_x;
        x = x/norm(x);
        lambda = x'*A*x;
        err = norm(x - prev_x);
        
        if ((++steps == max_steps) || (err < tol))
        
            return;
        
        endif
    
        prev_x = x;
    
    endwhile
    
endfunction
```

<!-- grade_id: cell-1851b004d61afd21, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - max(abs(eig(A)))) < 1e-4);

% Test 2
do
    A = rand(4);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - max(abs(eig(A)))) < 1e-4);

% Test 3
do
    A = rand(5);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - max(abs(eig(A)))) < 1e-4);

% Test 4
do
    A = rand(6);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - max(abs(eig(A)))) < 1e-4);
```

## Problema 2

Implementati funcția ```i_power_method(A, x0, tol, max_steps)``` ce primește o matrice pătratică $A$, o aproximare ințială a vectorului propriu cel mai puțin domninant $x_0$, o toleranță sub care se poate opri algoritmul $tol$ și un număr maxim de pași $max\_steps$ și returnează o valoarea proprie cel mai puțin dominantă a matricei $A$, $lambda$, un vector propriu normat asociat valorii proprii $lambda$, $x$, eroarea de la ultimul pas $err$ și numărul de pași efectuați $steps$.

Eroarea de la fiecare pas se calculează în funcție de vectorul propriu aproximat, nu de valoarea proprie aproximată.

Puteți folosi operatorul \\.

<!-- grade_id: cell-41090bd1c4ed1005, points: 0, grade: false, solution: true -->
```octave
function [lambda, x, err, steps] = i_power_method(A, x0, tol, max_steps)

    invA = inv(A);
    steps = 0;
    prev_x = x0;
    while 1
    
        x = invA*prev_x;
        x = x/norm(x);
        lambda = x'*A*x;
        err = norm(x - prev_x);
        
        if ((++steps == max_steps) || (err < tol))
        
            return;
        
        endif
    
        prev_x = x;
    
    endwhile

endfunction
```

<!-- grade_id: cell-b1c98ad65494eff0, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = i_power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - min(abs(eig(A)))) < 1e-4);

% Test 2
do
    A = rand(4);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = i_power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - min(abs(eig(A)))) < 1e-4);

% Test 3
do
    A = rand(5);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = i_power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - min(abs(eig(A)))) < 1e-4);

% Test 4
do
    A = rand(6);
until max(abs(imag(eig(A)))) < eps;
[lambda, x, err, steps] = i_power_method(A, rand(length(A), 1), 1e-8, 100);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(abs(lambda) - min(abs(eig(A)))) < 1e-4);
```

## Problema 3

Implementati funcția ```s_i_power_method(A, x0, tol, max_steps)``` ce primește o matrice pătratică $A$, o aproximare ințială a unui vector propriu $x_0$, o toleranță sub care se poate opri algoritmul $tol$, un număr maxim de pași $max\_steps$ și o valoare $alfa$ și returnează o valoarea proprie a matricei $A$ care este cea mai apropiată de $alfa$, $lambda$, un vector propriu normat asociat valorii proprii $lambda$, $x$, eroarea de la ultimul pas $err$ și numărul de pași efectuați $steps$.

Eroarea de la fiecare pas se calculează în funcție de vectorul propriu aproximat, nu de valoarea proprie aproximată.

Puteți folosi operatorul \\.

Nu se va folosi câtul Rayleigh întrucât ajungând prea repede la soluție, este greu de evaluat, matricile devenind singulare.

<!-- grade_id: cell-68966d23ba29ce42, points: 0, grade: false, solution: true -->
```octave
function [lambda, x, err, steps] = s_i_power_method(A, x0, tol, max_steps, alfa)
    invA = inv(A-alfa*eye(length(A)));
    steps = 0;
    prev_x = x0;
    while 1
    
        x = invA*prev_x;
        x = x/norm(x);
        lambda = x'*A*x;
        err = norm(x - prev_x);
        
        if ((++steps == max_steps) || (err < tol))
        
            return;
        
        endif
    
        prev_x = x;
    
    endwhile

endfunction
```

<!-- grade_id: cell-b47151dd04879aeb, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3);
until max(abs(imag(eig(A)))) < eps;
alfa_ind = randi(length(eig(A)));
alfa = eig(A)(alfa_ind) + rand() * 1e-1;
[lambda, x, err, steps] = s_i_power_method(A, rand(length(A), 1), 1e-8, 500, alfa);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(lambda - eig(A)(alfa_ind)) < 1e-4);
assert(norm(norm(x) - 1) < 1e-4);

% Test 2
do
    A = rand(4);
until max(abs(imag(eig(A)))) < eps;
alfa_ind = randi(length(eig(A)));
alfa = eig(A)(alfa_ind) + rand() * 1e-1;;
[lambda, x, err, steps] = s_i_power_method(A, rand(length(A), 1), 1e-8, 500, alfa);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(lambda - eig(A)(alfa_ind)) < 1e-4);
assert(norm(norm(x) - 1) < 1e-4);

% Test 3
do
    A = rand(5);
until max(abs(imag(eig(A)))) < eps;
alfa_ind = randi(length(eig(A)));
alfa = eig(A)(alfa_ind) + rand() * 1e-1;;
[lambda, x, err, steps] = s_i_power_method(A, rand(length(A), 1), 1e-8, 100, alfa);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(lambda - eig(A)(alfa_ind)) < 1e-4);
assert(norm(norm(x) - 1) < 1e-4);

% Test 4
do
    A = rand(6);
until max(abs(imag(eig(A)))) < eps;
alfa_ind = randi(length(eig(A)));
alfa = eig(A)(alfa_ind) + rand() * 1e-1;;
[lambda, x, err, steps] = s_i_power_method(A, rand(length(A), 1), 1e-8, 100, alfa);
assert(norm(lambda * x - A * x) < 1e-4);
assert(norm(lambda - eig(A)(alfa_ind)) < 1e-4);
assert(norm(norm(x) - 1) < 1e-4);
```

## Problema 4

Implementati funcția ```orthogonal_power_method(A, tol, max_steps)``` ce primește o matrice simetrică $A$, o toleranță sub care se poate opri algoritmul $tol$ și un număr maxim de pași pentru fiecare metodă a puterii aplicată, $max\_steps$, și returnează o matrice $X$ care are pe fiecare coloană câte un vector propriu normat (de la cel mai dominant la cel mai puțin dominant), un vector linie $l$ ce conține toate valorile proprii (de la cea mai dominantă la aceea cel mai puțin dominantă) și $total\_steps$, numărul total de pași (suma pașilor făcuți pentru fiecare metodă a puterii).

Eroarea de la fiecare pas se calculează în funcție de vectorul propriu aproximat, nu de valoarea proprie aproximată.

<!-- grade_id: cell-7074a1b524b8b51c, points: 0, grade: false, solution: true -->
```octave
function [X, l, total_steps] = orthogonal_power_method(A, tol, max_steps)
    
    X = [];
    l = [];
    n = length(A);

    for i = 1:n
    
        prev_x = rand(n, 1);
        total_steps = 0;
        
        while 1
        
            for j = 1:size(X, 2)
            
                prev_x = prev_x - X(:, j)' * prev_x * X(:, j);
            
            endfor
            
            x = A*prev_x;
            x = x/norm(x);
            
            lambda = x'*A*x;
            err = norm(x - prev_x);
            
            if ((++total_steps == max_steps) || (err < tol))
            
                l = [l, lambda];
                X = [X, x];
                break;
            
            endif
        
            prev_x = x;
        
        endwhile
    
    endfor

endfunction
```

<!-- grade_id: cell-1db5f0490bcef817, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_power_method(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 2
do
    A = rand(4); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_power_method(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 3
do
    A = rand(5); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_power_method(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 4
do
    A = rand(6); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_power_method(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);
```

## Problema 5

Implementati funcția ```orthogonal_iteration(A, tol, max_steps)``` ce primește o matrice simetrică $A$, o toleranță sub care se poate opri algoritmul $tol$ și un număr maxim de pași, $max\_steps$, și returnează o matrice $X$ care are pe fiecare coloană câte un vector propriu normat (de la cel mai dominant la cel mai puțin dominant), un vector linie $l$ ce conține toate valorile proprii (de la cea mai dominantă la aceea cel mai puțin dominantă) și $total\_steps$, numărul total de pași (suma pașilor făcuți pentru fiecare metodă a puterii).

Eroarea de la fiecare pas se calculează în funcție de matricea X aproximată.

<!-- grade_id: cell-569f040cca94bb06, points: 0, grade: false, solution: true -->
```octave
function [X, l, steps] = orthogonal_iteration(A, tol, max_steps)
 
     n = length(A);
     U = rand(n);
     prev_U = U;
     
     X = [];
     steps = 0;
     
     while 1
     
         [Q, R] = qr(U);
         U = A*Q;
     
        if ((++steps == max_steps) || (norm(U - prev_U) < tol))
        
                X = [X, Q];
                [Q, R] = qr(U);
                l = diag(Q'*A*Q)';
                
            return;
        
        endif
     
         prev_U = U;
     
     endwhile
 

    
endfunction
```

<!-- grade_id: cell-38b60813509ab6c0, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_iteration(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 2
do
    A = rand(4); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_iteration(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 3
do
    A = rand(5); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_iteration(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 4
do
    A = rand(6); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = orthogonal_iteration(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);
```

## Problema 6

Implementati funcția ```deflation(A, tol, max_steps)``` ce primește o matrice $A$, o toleranță sub care se pot opri MPD și MPID, $tol$, și un număr maxim de pași pentru MPD și MPID, $max\_steps$, și returnează o matrice $X$ care are pe fiecare coloană câte un vector propriu normat (de la cel mai dominant la cel mai puțin dominant), un vector linie $l$ ce conține toate valorile proprii (de la cea mai dominantă la aceea cel mai puțin dominantă) și $total\_steps$, numărul total de pași (suma pașilor făcuți pentru fiecare metodă a puterii).

MPD = Metoda Puterii Directe

MPID = Metoda Puterii Inverse Deplasată

<!-- grade_id: cell-b7d49a670624b96b, points: 0, grade: false, solution: true -->
```octave
function [X, l, total_steps] = deflation(A, tol, max_steps)

    X=[];
    l=[];
    A0 = A;
    n = length(A);
    total_steps = 0;
    
    for i = 1:n
    
        [lambda, xb, err, steps] = power_method(A, rand(length(A), 1), tol, max_steps);
        
        total_steps = total_steps + steps;
        
        [lambda, xi, err, steps] = s_i_power_method(A0, rand(n, 1), tol, max_steps, lambda + rand()*1e-2);
        total_steps = total_steps + steps;
        
        X = [X, xi];
        l = [l, lambda];
        
        if (i == n)
        
            return;
        
        endif
        
        v = xb;
        v(1) = v(1) - sign(xb(1))*norm(xb);
        H = eye(length(A)) - 2*(v * v')/(v' * v);
    
        A = H'*A*H;
        A = A(2:end, 2:end);
    
    endfor



endfunction
```

<!-- grade_id: cell-a778cccc183a1eb4, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3);
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = deflation(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
for i = 1:length(A)
    assert(norm(norm(X(:, i)) - 1) < 1e-4);
endfor

% Test 2
do
    A = rand(4);
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = deflation(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
for i = 1:length(A)
    assert(norm(norm(X(:, i)) - 1) < 1e-4);
endfor

% Test 3
do
    A = rand(5);
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = deflation(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
for i = 1:length(A)
    assert(norm(norm(X(:, i)) - 1) < 1e-4);
endfor

% Test 4
do
    A = rand(6);
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = deflation(A, 1e-8, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
for i = 1:length(A)
    assert(norm(norm(X(:, i)) - 1) < 1e-4);
endfor
```

## Problema 7

Implementati funcția ```hessenberg(A)``` ce primește o matrice $A$ și returnează o matrice tridiagonală cu aceleași valori proprii ca $A$ dacă $A$ este simetrică sau o matrice superior Hessenberg ce aceleași valori proprii ca $A$ dacă a nu este simetrică.

HINT: Trebuie implementat fix același algoritm pentru ambele cazuri, nimic diferit în cod.

<!-- grade_id: cell-7b94e65947b636f3, points: 0, grade: false, solution: true -->
```octave
function [R_] = hessenberg(A)

    R_ = 1;
    n = length(A);

        
    for k = 1:n-2

        w = zeros(n, 1);
        alpha = -sign(A(k+1, k))*norm(A(k+1:n, k));
        r = sqrt(1/2*alpha^2 - 1/2*alpha*A(k+1, k));

        w(k+1) = (A(k+1, k) - alpha)/2/r;

        for j = k+2:n

            w(j) = A(j, k)/2/r;

        endfor

        P = eye(n) - 2*w*w';

        A = P*A*P;

    endfor
    
    R_ = A;

end
```

<!-- grade_id: cell-8ab7c7070dfde037, points: 3, grade: true, solution: false -->
```octave
% Test 1
A = rand(3);
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);

% Test 2
A = rand(5);
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);

% Test 3
A = rand(7);
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);

% Test 4
A = rand(20);
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);

% Test 5
A = rand(3); A = A * A';
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);
assert(max(abs(triu(R_, 2))) < 1e-4);

% Test 6
A = rand(5); A = A * A';
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);
assert(max(abs(triu(R_, 2))) < 1e-4);

% Test 7
A = rand(7); A = A * A';
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);
assert(max(abs(triu(R_, 2))) < 1e-4);

% Test 8
A = rand(20); A = A * A';
R_ = hessenberg(A);
assert(norm(sort(eig(A)) - sort(eig(R_))) < 1e-4);
assert(max(abs(tril(R_, -2))) < 1e-4);
assert(max(abs(triu(R_, 2))) < 1e-4);
```

## Problema 8

Implementati funcția ```qr_algorithm(A, max_steps)``` ce primește o matrice $A$ simetrică și un număr maxim de pași, $max\_steps$, și returnează o matrice $X$ care are pe fiecare coloană câte un vector propriu normat (de la cel mai dominant la cel mai puțin dominant), un vector linie $l$ ce conține toate valorile proprii (de la cea mai dominantă la aceea cel mai puțin dominantă) și $total_steps$, numărul total de pași (suma pașilor făcuți pentru fiecare metodă a puterii).

HINT: $X = Q_1 * Q_2 * ...$

<!-- grade_id: cell-ecb82f560c6007b1, points: 0, grade: false, solution: true -->
```octave
function [X, l, steps] = qr_algorithm(A, max_steps)

    X = [];
    l = [];
    steps = 0;
    prev_A = A;
    X = eye(length(A));
    
   while 1
    
        [Q, R] = qr(prev_A);
        X = X*Q;
        A = R * Q;

        if ((++steps == max_steps) || (diag(A) == diag(prev_A)))

            l = diag(A)';
            return;


        endif

        prev_A = A;

    endwhile

endfunction
```

<!-- grade_id: cell-749288febcfb84a8, points: 3, grade: true, solution: false -->
```octave
% Test 1
do
    A = rand(3); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = qr_algorithm(A, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 2
do
    A = rand(4); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = qr_algorithm(A, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 3
do
    A = rand(5); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = qr_algorithm(A, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);

% Test 4
do
    A = rand(6); A = A * A';
until max(abs(imag(eig(A)))) < eps;
[X, l, steps] = qr_algorithm(A, 500);
assert(norm(A * X - X * diag(l)) < 1e-4);
assert(norm(sort(l') - sort(eig(A))) < 1e-4);
assert(norm(X) - 1 < 1e-4);
```

```octave
% QR FARA DEPLASARE LAB_MN_08_CERINTA_1
```

```octave
function [P] = getMatrRotatie(A, i ,j, n)

    % il facem pe A(i, j) = 0
    % sn = sin
    % cs = cos
    sn = -A(i, j)/sqrt(A(j, j)^2 + A(i, j)^2);
    cs = A(j, j)/sqrt(A(j, j)^2 + A(i, j)^2);
    
    P = eye(n);
    
    P(j, j) =  P(i, i)= cs;
    P(i, j) = sn;
    P(j, i) = -sn;
    

endfunction
```

```octave
A = [1441 12 0 0 0; 12 293 943 0 0; 0 943 451 9 0; 0 0 9 41 2; 0 0 0 2 131];
valPropriiReale = eig(A);

% dimensiunea matricei A
n = length(A);

valPropriiAprox = [];

% incepem de la P2 (coresp (2, 1))


for i = n:-1:2

    while (abs(A(i, i-1)) > 1e-3)

        j = 1;
        Q = eye(i);
        for k = 2:i

                P = getMatrRotatie(A, k, j, i);
                A = P * A;
                Q = Q * P';


                j++;

        endfor

        
        R = A;

        A = R * Q;

    endwhile
    
    A;

    valPropriiAprox = [valPropriiAprox, A(i, i)];

    A = A(1:i-1, 1:i-1);
    

endfor

% in final ramanem doar cu A(1, 1)
valPropriiAprox = [valPropriiAprox, A(1, 1)];

disp(valPropriiAprox);
disp(valPropriiReale');
norm(sort(valPropriiAprox)-sort(valPropriiReale'))
```

Output:
```
40.982    131.044   -574.402   1317.802   1441.574
   -574.402     40.982    131.044   1317.802   1441.574
ans =    1.0150e-08
```

```octave
% QR CU DEPLASARE LAB_MN_08_CERINTA_1
```

```octave
function [minSigma] = dmin(a, b, pct)

    minValue = a-pct;
    minSigma = a;
    if (abs(b-pct) < abs(minValue))

        minSigma = b;
        minValue = b-pct;
    
    endif

endfunction

```

```octave
A = [3 1 0; 1 3 1; 0 1 3];

valPropriiReale = eig(A);

% dimensiunea matricei A
n = length(A);

valPropriiAprox = [];

sigmauri = 0;

for i = n:-1:3

    A = A(1:i, 1:i);

    val = 0;

    while (abs(A(i, i-1)) > 1e-8)
    
        valprop = eig(A(i-1:i, i-1:i));
        sigma = dmin(valprop(1), valprop(2), A(i, i));
        
        val = val+sigma;
    
        A = A - sigma * eye(i);

        j = 1;
        Q = eye(i);
        for k = 2:i

                P = getMatrRotatie(A, k, j, i);
                A = P * A;
                Q = Q * P';

                j++;

        endfor

        
        R = A;

        A = R * Q;


    endwhile

    break;
    val = val + A(i, i) + sigmauri;
    
    sigmauri = val - A(i, i);

    valPropriiAprox = [valPropriiAprox, val];

    break;

endfor
A

A = A(1:2, 1:2);

valPropriiAprox = [valPropriiAprox, sigmauri + eig(A)'];

disp(sort(valPropriiAprox));
disp(sort(valPropriiReale'));
norm(sort(valPropriiAprox)-sort(valPropriiReale'))
```

Output:
```
A =

   2.00000   0.70711  -0.00000
   0.70711   1.00000  -0.70711
   0.00000  -0.70711   0.00000

   0.63397   2.36603
   1.5858   3.0000   4.4142
error: operator -: nonconformant arguments (op1 is 1x2, op2 is 1x3)
```

```octave

```

```octave

```